# Программирование для всех <br>(основы работы с Python)

*Разработчик формата сдачи заданий, фабрики автоматических тестов, автор задачи 7: Щуров И.В.<br>
Автор задач: Тамбовцева А.А.*

**За разные задачи можно получить разное число баллов. Чтобы получить оценку 10, достаточно набрать 10 баллов. Вы можете решить больше задач, чем требуется, чтобы потренироваться.**

Задача считается решенной, если код успешно проходит автоматические тесты и не содержит грубых нарушений («подгонка» под тесты, копирование кода другого студента). «Подгонка» под тесты – написание кода не для общего случая, а для соответствия конкретным значениям в ячейке с тестами. Пример: пишем не универсальный код для проверки четности числа, а код, который правильно работает только на числах 2, 5, 8, потому что именно эти числа используются в `test_data` .

### Формат сдачи

Для сдачи домашнего задания необходимо скачать текущий ipynb-файл, вписать в него свои решения, сохранить изменения и загрузить обновленный файл на Dropbox по предложенной на странице курса ссылке. Иметь аккаунт Dropbox для этого не требуется, достаточно пройти по ссылке и выбрать файл на компьютере.

Получить ipynb-файл можно, выбрав в Jupyter пункт меню *File → Download as... → Notebook (.ipynb)*.

### Тесты для проверки кода

На первом этапе проверки ваши решения в домашних заданиях и лабораторных работах будут проверяться с помощью автоматических тестов – ячеек с кодом, которые проверяют соответствие результатов исполнения ваших программ ожидаемым результатам в условиях задач. Эти тесты можно использовать для проверки кода перед сдачей задания.

Итак, для предварительной проверки задания нужно сделать следующее:

1. Скачать данный ipynb-файл на свой компьютер, открыть его в Jupyter Notebook.
2. Активировать тесты (см. ниже).
3. Вставить решение каждой задачи в ячейку для кода, следующую за его условием, там, где написан комментарий `# YOUR CODE HERE`. Отступ вашего кода должен составлять 4 пробела. Ничего не менять вокруг!
4. Запустить ячейку, в которую вы вставили код с решением. Ввести какие-то входные данные, проверить визуально правильность вывода.
5. Запустить следующую ячейку (в ней содержится тест). Если запуск ячейки с тестом не приводит к появлению ошибки (`assertion`), значит, всё в порядке, задача решена. Если приводит к появлению ошибки, значит, тест не пройден и нужно искать ошибку.

Внимание! Если в какой-то момент забыть ввести входные данные и перейти на следующую ячейку, есть риск, что Notebook перестанет откликаться. В этом случае надо перезапустить ядро: *Kernel → Restart*. При этом потеряются все значения переменных, но сам код останется.

### Активировать тесты

Запустите следующую ячейку, чтобы иметь возможность запускать тесты.

```python
# Фабрика тестов для проверки программ, принимающих данные через input()

from collections import deque

class Tester(object):
    def __init__(self, inp):
        self.outputs = []
        self.inputs = deque(inp)
    def print(self, *args, sep = " ", end = "\n"):
        text = sep.join(map(str, args)) + end
        newlines = text.splitlines(keepends=True)
        if self.outputs and self.outputs[-1] and self.outputs[-1][-1] != "\n" and newlines:
            self.outputs[-1] += newlines[0]
            self.outputs.extend(newlines[1:])
        else:
            self.outputs.extend(newlines)

    def input(self, *args):
        assert self.inputs, "Вы пытаетесь считать больше элементов, чем предусмотрено условием"
        return self.inputs.popleft()
    def __enter__(self):
        global print
        global input
        print = self.print
        input = self.input
        return self.outputs
    def __exit__(self, *args):
        global print
        global input
        del print
        del input
```

### Задача 1. На дальней станции сойду... (1 балл)

Дан список станций желтой ветки московского метро (упорядочены как на карте):

    line = ["Третьяковская", "Марксистская", "Площадь Ильича", "Авиамоторная", "Шоссе Энтузиастов", "Перово", "Новогиреево", "Новокосино"]

Давайте считать, что движение происходит от Третьяковской до Новокосино.

Напишите программу, которая запрашивает у пользователя название текущей станции метро (без лишнего пробела на конце) и выводит на экран сообщение вида

    Следующая станция: [станция].

Считайте, что пользователь знает, какая станция конечная, и не будет запрашивать станцию, следующую после Новокосино.

**Пример**

Входные данные:

    Площадь Ильича

Выходные данные:

    Следующая станция: Авиамоторная.

<!-- grade_id: ans01, points: 0, grade: false, solution: true -->
```python
def metro():
    line = ["Третьяковская", "Марксистская", "Площадь Ильича", "Авиамоторная", "Шоссе Энтузиастов",
    "Перово", "Новогиреево", "Новокосино"]

    # YOUR CODE HERE
    a = str(input())
    d = line[line.index(a)+1]
    print('Следующая станция: ',d,'.',sep='')


metro()
```

Output:
```
Третьяковская
Следующая станция: Марксистская.
```

<!-- grade_id: test01, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("Третьяковская", "Следующая станция: Марксистская."),
    ("Марксистская", "Следующая станция: Площадь Ильича."),
    ("Площадь Ильича", "Следующая станция: Авиамоторная."),
    ("Авиамоторная", "Следующая станция: Шоссе Энтузиастов."),
    ("Шоссе Энтузиастов","Следующая станция: Перово."),
    ("Перово", "Следующая станция: Новогиреево."),
    ("Новогиреево", "Следующая станция: Новокосино.")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        metro()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 2. Любопытство – не порок (1 балл)

Любопытный питон Петя ловит прохожих на улице и просит назвать их какую-нибудь цифру от 0 до 9. Все ответы он записывает в виде последовательности из цифр, а затем считает сумму этих цифр. Итак, напишите программу, которая:

* сначала запрашивает число прохожих с клавиатуры (число $n$);
* затем $n$ раз просит ввести какую-нибудь цифру;
* выводит на экран полученную последовательность цифр и их сумму через пробел.

Считайте, что пользователь всегда вводит с клавиатуры корректные значения, целые положительные числа для количества прохожих и цифры от 0 до 9.

Пользоваться готовой функцией `sum()` для вычисления суммы элементов списка нельзя.

**Примеры работы программы**

Ввод:

    4
    0
    9
    8
    2

Вывод:

    0982 19
    
Ввод:

    3
    1
    1
    5

Вывод:

    115 7
    
Ввод:

    7
    9
    9
    8
    1
    2
    3
    0

Вывод:

    9981230 32

<!-- grade_id: ans02, points: 0, grade: false, solution: true -->
```python
def curious():

    # YOUR CODE HERE
    a = int(input())
    i=0
    j=0
    c=[]
    while i < a:
      b=int(input())
      c.append(str(b))
      j+=b
      i+=1
    print(''.join(c),j)

curious()
```

Output:
```
3
1
1
5
115 7
```

<!-- grade_id: test02, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("4 5 6 0 1", "5601 12"),
    ("7 9 9 8 1 2 3 0", "9981230 32"),
    ("2 0 1", "01 1"),
    ("3 1 1 5", "115 7"),
    ("10 3 4 7 7 3 1 0 2 3 5", "3477310235 35"),
    ("1 0", "0 0")
]

for inp, out in test_data:
    with Tester(inp.split()) as t:
        curious()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0].strip() == str(out), "Неверный ответ, были введены числа " + inp
```

### Задача 3. Чем больше, тем веселее (1 балл)

В массиве `valid` сохранено число действительных бюллетеней на 5 избирательных участках, а в массиве `invalid` – число недействительных бюллетеней на тех же участках. Также известно общее число избирателей на каждом участке, значения сохранены в массиве `total`.

```python
import numpy as np

valid = np.array([450, 800, 630, 740, 950])
invalid = np.array([10, 200, 100, 50, 120])
total = np.array([900, 1000, 1200, 890, 2300])
```

Используя имеющиеся массивы, вычислите явку (в процентах) на каждом избирательном участке, сохраните результаты в массив `turnout` и выведите на экран с новой строки его элементы, округлённые до второго знака после запятой.

*Напоминание.* Явка на каждом участке определяется как сумма действительных и недействительных бюллетеней, делённая на общее число зарегистированных избирателей (не забудьте перевести полученную долю в проценты).

**Примеры работы программы**

На выходе:

    51.11
    96.02
    60.83
    88.76
    46.52

<!-- grade_id: ans04, points: 0, grade: false, solution: true -->
```python
def elections():

    valid = np.array([450, 787, 630, 740, 950])
    invalid = np.array([10, 202, 100, 50, 120])
    total = np.array([900, 1030, 1200, 890, 2300])

    # YOUR CODE HERE
    c = ((valid + invalid)/total)*100
    print(*np.around(c,2), sep ='\n')
elections()
```

Output:
```
51.11
96.02
60.83
88.76
46.52
```

<!-- grade_id: test04, points: 1, grade: true, solution: false -->
```python
test_data = [("", ["51.11\n", "96.02\n", "60.83\n",
                   "88.76\n", "46.52\n"])]

for inp, out in test_data:
    with Tester(inp.split()) as t:
        elections()
        line_t = "".join(t)
        t = line_t.splitlines()
        assert len(t) == len(out), "Неверное количество строк в выводе"
        for l_test, l_out in zip(t, out):
            assert len(l_test.split()) == len(l_out.split()), \
                "Неверное количество элементов в строке " + l_out
            for el_test, el_out in zip(l_test.split(), l_out.split()):
                assert el_test == el_out, "Ошибка {} != {}".format(l_test,
                                                                   l_out)
```

### Задача 4. Жизнь летит, как ураган... (2 балла)

В городе Дакбург намечаются выборы. Напишите программу, которая запрашивает имена кандидатов через запятую, считает их количество и выводит на экран число способов упорядочить их в избирательном бюллетене.

*Подсказка.* Число способов упорядочить $n$ объектов вычисляется с помощью факториала. Факториалом $n!$ натурального числа $n$ называется произведение всех натуральных чисел от 1 до $n$ включительно. Пример:

$$
4! = 1 \times 2 \times 3 \times 4 = 24
$$

Использовать при решении какие-либо готовые библиотечные функции нельзя (то есть запрещена конструкция `import`).

Ввод:

    Билли,Вилли,Дилли

Вывод:

    6
    
Ввод:

    Скрудж Макдак,Флинтхарт Гломгольд

Вывод:

    2
    
Ввод:

    Зигзаг,Дональд,Клювдия,Винт

Вывод:

    24

<!-- grade_id: ans05, points: 0, grade: false, solution: true -->
```python
def my_fact():

    # YOUR CODE HERE
    a = str(input())
    i=0
    s=1
    for k in a.split(','):
      i+=1

    for j in range(1,i+1):
      s*=j
    print(s)
my_fact()
```

Output:
```
рпрпр,опооп,прпр
6
```

<!-- grade_id: test05, points: 2, grade: true, solution: false -->
```python
test_data = [
    ("Скрудж Макдак,Флинтхарт Гломгольд", 2),
    ("Билли,Вилли,Дилли", 6),
    ("Зигзаг,Дональд,Клювдия,Винт", 24),
    ("A,B,C,D,E", 120),
    ("0,1,2,3,4,5,6,7,8,9", 3628800)
]

for inp, out in test_data:
    with Tester([inp]) as t:
        my_fact()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0].strip() == str(out), "Неверный ответ, были введены числа " + inp
```

### Задача 5. Кортеж Его Высочества (2 балла)

Чиполлино попросил медведя, освобождённого им из зоопарка, собрать некоторую информацию о кортеже принца Лимона, который должен проехать вечером по центральной улице, но забыл, что тот не знает, как пишутся цифры. Тогда Чиполлино придумал такую систему обозначений:

* `&` – принц Лимон (носит плащ с длинным «хвостом»);
* `8` – министры (носят тугие ленты поверх мундира);
* `0` – придворные лимончики (без лент и плащей).

Так как порядок следования в кортеже строго регламентирован, на первом месте всегда следует принц, далее – министры, а затем уже придворные лимончики. Исходя из этого, у медведя могут получаться записи такого вида: `& 88 0000` (принц, два министра, четверо придворных), `& 8 00000`(принц, министр, пятеро придворных), и так далее.

Напишите программу, которая запрашивает с клавиатуры строку, состоящую из символов, описанных выше, и пробелов, и возвращает кортеж из этих символов и чисел, которые соответствуют количеству лимонов, занимающих разные должности.

Ввод:

    & 8 00000

Вывод:

    ('&', 1, '8', 1, '0', 5)
    
Ввод:

    & 88 0000

Вывод:

    ('&', 1, '8', 2, '0', 4)
    
Ввод:

    & 888 00

Вывод:

    ('&', 1, '8', 3, '0', 2)

<!-- grade_id: ans06, points: 0, grade: false, solution: true -->
```python
def lemongrass():

    # YOUR CODE HERE
    a = input()
    k=[]
    for i in a:
      k.append(i)
    b = k.count('&')
    c = k.count('8')
    d = k.count('0')
    my_tuple = ('&',b,'8',c,'0',d)
    print(my_tuple)

lemongrass()
```

Output:
```
& 888 00
('&', 1, '8', 3, '0', 2)
```

<!-- grade_id: test06, points: 2, grade: true, solution: false -->
```python
test_data = [
    ("& 888 0", "('&', 1, '8', 3, '0', 1)"),
    ("& 8 00000", "('&', 1, '8', 1, '0', 5)"),
    ("& 8 000", "('&', 1, '8', 1, '0', 3)"),
    ("& 8888 000000", "('&', 1, '8', 4, '0', 6)"),
    ("& 88 000000", "('&', 1, '8', 2, '0', 6)")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        lemongrass()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0].strip() == str(out), "Неверный ответ, были введены числа " + inp
```

### Задача 6. Сладкая жизнь (3 балла)

Реализуйте программу, которая фильтрует угощения из «Сладкого королевства» на основе их стоимости.

* С клавиатуры запрашивается перечень названий сладостей (строки, разделены запятыми без пробела).
* С клавиатуры запрашивается перечень цен сладостей (целые числа, разделены запятыми без пробела). Каждая цена соответствует товару из предыдущего перечня, то есть гарантируется одинаковое число товаров и цен.
* Гарантируется, что будет введена информация минимум для одного товара.
* С клавиатуры запрашивается строка — минимальная и максимальная цена, которую ученики Хогвартса готовы заплатить за сладость (через пробел).
* Программа выводит названия сладостей, цены на которые не менее минимального введенного числа и не более максимального введенного числа.

**Примеры работы программы:**

Ввод:

    шоколадная лягушка,перечный чертик
    3,4
    0 3

Вывод:

    шоколадная лягушка
    
    
Ввод:

    тараканьи грозди,берти ботс,шоколадный котелок,сахарное перо
    8,6,4,3
    4 7

Вывод:

    берти ботс
    шоколадный котелок
    
Ввод:

    тыквенное печенье,сахарная мышка,берти ботс
    3,3,6
    1 2
    
Вывод:

    [ничего не выводится]

```python
def problem_hd():
    s = []
    a = input().split(',')
    b = input().split(',')
    min, max = map(int, input().split())

    for i in range(len(b)):
      if min <= int(b[i]) and int(b[i]) <= max:
        s.append(a[i])
    for j in s:
      print(j)

problem_hd()
```

Output:
```
сок,мок
3,4
0 3
сок
```

```python
test_data = [("шоколадная лягушка,перечный чертик;3,4;0 3",
              ["шоколадная лягушка"]),
             ("тараканьи грозди,берти ботс,шоколадный котелок,сахарное перо;8,6,4,3;4 7",
              ["берти ботс", "шоколадный котелок"]),
             ("тыквенное печенье,сахарная мышка,берти ботс;3,3,6;1 2", []),
             ("кровяной леденец,мятный леденец,помадка;8,2,3;6 10", ["кровяной леденец"]),
             ("тараканьи грозди,берти ботс,шоколадный котелок,сахарное перо,шоколад;8,6,3,4,3;0 2", [])
             ]
for inp, out in test_data:
    with Tester(inp.split(";")) as t:
        problem_hd()
        line_t = "".join(t)
        t = line_t.splitlines()
        assert len(t) == len(out), "Неверный ответ"
        for l_test, l_out in zip(t, out):
            assert len(l_test.split()) == len(l_out.split()), "Неверный ответ"
            for el_test, el_out in zip(l_test.split(), l_out.split()):
                assert el_test == el_out, "Неверный ответ"
```

### Задача 7.  Таблица умножения (3 балла)

Введите число $n$ с клавиатуры. Выведите таблицу умножения от 1 до $n$ включительно, как показано в примере.

**Пример работы программы**

Ввод:

    5

Вывод:

    1 2 3 4 5
    2 4 6 8 10
    3 6 9 12 15
    4 8 12 16 20
    5 10 15 20 25

<!-- grade_id: ans07, points: 0, grade: false, solution: true -->
```python
def mult_table():

    # YOUR CODE HERE
    a = int(input())
    empty=[]
    for i in range(1,a+1):
      for j in range(1,a+1):
        empty.append(i*j)
    for i in range(0, len(empty), a):
      print(*empty[i:i+a])
mult_table()
```

Output:
```
5
1 2 3 4 5
2 4 6 8 10
3 6 9 12 15
4 8 12 16 20
5 10 15 20 25
```

<!-- grade_id: test07, points: 3, grade: true, solution: false -->
```python
test_data = [(1, ['1 \n']),
 (4, ['1 2 3 4 \n', '2 4 6 8 \n', '3 6 9 12 \n', '4 8 12 16 \n']),
 (7,
  ['1 2 3 4 5 6 7 \n',
   '2 4 6 8 10 12 14 \n',
   '3 6 9 12 15 18 21 \n',
   '4 8 12 16 20 24 28 \n',
   '5 10 15 20 25 30 35 \n',
   '6 12 18 24 30 36 42 \n',
   '7 14 21 28 35 42 49 \n']),
 (10,
  ['1 2 3 4 5 6 7 8 9 10 \n',
   '2 4 6 8 10 12 14 16 18 20 \n',
   '3 6 9 12 15 18 21 24 27 30 \n',
   '4 8 12 16 20 24 28 32 36 40 \n',
   '5 10 15 20 25 30 35 40 45 50 \n',
   '6 12 18 24 30 36 42 48 54 60 \n',
   '7 14 21 28 35 42 49 56 63 70 \n',
   '8 16 24 32 40 48 56 64 72 80 \n',
   '9 18 27 36 45 54 63 72 81 90 \n',
   '10 20 30 40 50 60 70 80 90 100 \n']),
 (13,
  ['1 2 3 4 5 6 7 8 9 10 11 12 13 \n',
   '2 4 6 8 10 12 14 16 18 20 22 24 26 \n',
   '3 6 9 12 15 18 21 24 27 30 33 36 39 \n',
   '4 8 12 16 20 24 28 32 36 40 44 48 52 \n',
   '5 10 15 20 25 30 35 40 45 50 55 60 65 \n',
   '6 12 18 24 30 36 42 48 54 60 66 72 78 \n',
   '7 14 21 28 35 42 49 56 63 70 77 84 91 \n',
   '8 16 24 32 40 48 56 64 72 80 88 96 104 \n',
   '9 18 27 36 45 54 63 72 81 90 99 108 117 \n',
   '10 20 30 40 50 60 70 80 90 100 110 120 130 \n',
   '11 22 33 44 55 66 77 88 99 110 121 132 143 \n',
   '12 24 36 48 60 72 84 96 108 120 132 144 156 \n',
   '13 26 39 52 65 78 91 104 117 130 143 156 169 \n'])]

for inp, out in test_data:
    with Tester([str(inp)]) as t:
        mult_table()
        line_t = "".join(t)
        t = line_t.splitlines()
        assert len(t) == len(out), "Неверное количество строк в выводе"
        for l_test, l_out in zip(t, out):
            assert len(l_test.split()) == len(l_out.split()), \
                "Неверное количество элементов в строке " + l_out
            for el_test, el_out in zip(l_test.split(), l_out.split()):
                assert el_test == el_out, "Ошибка {} != {}".format(l_test,
                                                                   l_out)
print("Теперь вы можете печатать тетрадки для школьников!")
```

Output:
```
Теперь вы можете печатать тетрадки для школьников!
```