# Программирование для всех <br>(основы работы с Python)

*Разработчик формата сдачи заданий и фабрики автоматических тестов: Щуров И.В.<br>
Автор задач: Тамбовцева А.А.*

**За разные задачи можно получить разное число баллов. Чтобы получить оценку 10, достаточно набрать 8 баллов. Вы можете решить больше задач, чем требуется, чтобы потренироваться.**

Задача считается решенной, если код успешно проходит автоматические тесты и не содержит грубых нарушений («подгонка» под тесты, копирование кода другого студента). «Подгонка» под тесты – написание кода не для общего случая, а для соответствия конкретным значениям в ячейке с тестами. Пример: пишем не универсальный код для проверки четности числа, а код, который правильно работает только на числах 2, 5, 8, потому что именно эти числа используются в `test_data` .

**Условия задач повышенной сложности (на 2-3 балла) включают ограничения – запрет на использование некоторых конструкций и операторов. Это не искусственное усложнение жизни, это задачи, развивающие алгоритмическое мышление. Умение выполнить задачу с минимальными затратами по времени без лишнего кода с условиями и циклами крупные компании требуют сейчас даже от аналитиков. На собеседовании вполне может встретиться задача на алгоритмы, даже если вы не разработчик, и планируете мирно анализировать и визуализировать данные.**

### Формат сдачи

Для сдачи домашнего задания необходимо скачать текущий ipynb-файл, вписать в него свои решения, сохранить изменения и загрузить обновленный файл на Dropbox по предложенной на странице курса ссылке. Иметь аккаунт Dropbox для этого не требуется, достаточно пройти по ссылке и выбрать файл на компьютере.

Получить ipynb-файл можно, выбрав в Jupyter пункт меню *File → Download as... → Notebook (.ipynb)*.

### Тесты для проверки кода

На первом этапе проверки ваши решения в домашних заданиях и лабораторных работах будут проверяться с помощью автоматических тестов – ячеек с кодом, которые проверяют соответствие результатов исполнения ваших программ ожидаемым результатам в условиях задач. Эти тесты можно использовать для проверки кода перед сдачей задания.

Итак, для предварительной проверки задания нужно сделать следующее:

1. Скачать данный ipynb-файл на свой компьютер, открыть его в Jupyter Notebook.
2. Активировать тесты (см. ниже).
3. Вставить решение каждой задачи в ячейку для кода, следующую за его условием, там, где написан комментарий `# YOUR CODE HERE`. Отступ вашего кода должен составлять 4 пробела. Ничего не менять вокруг!
4. Запустить ячейку, в которую вы вставили код с решением. Ввести какие-то входные данные, проверить визуально правильность вывода.
5. Запустить следующую ячейку (в ней содержится тест). Если запуск ячейки с тестом не приводит к появлению ошибки (`assertion`), значит, всё в порядке, задача решена. Если приводит к появлению ошибки, значит, тест не пройден и нужно искать ошибку.

Внимание! Если в какой-то момент забыть ввести входные данные и перейти на следующую ячейку, есть риск, что Notebook перестанет откликаться. В этом случае надо перезапустить ядро: *Kernel → Restart*. При этом потеряются все значения переменных, но сам код останется.

### Активировать тесты

Запустите следующую ячейку, чтобы иметь возможность запускать тесты.

```python
# Фабрика тестов для проверки программ, принимающих данные через input()

from collections import deque

class Tester(object):
    def __init__(self, inp):
        self.outputs = []
        self.inputs = deque(inp)
    def print(self, *args, sep = " ", end = "\n"):
        text = sep.join(map(str, args)) + end
        newlines = text.splitlines(keepends=True)
        if self.outputs and self.outputs[-1] and self.outputs[-1][-1] != "\n" and newlines:
            self.outputs[-1] += newlines[0]
            self.outputs.extend(newlines[1:])
        else:
            self.outputs.extend(newlines)

    def input(self, *args):
        assert self.inputs, "Вы пытаетесь считать больше элементов, чем предусмотрено условием"
        return self.inputs.popleft()
    def __enter__(self):
        global print
        global input
        print = self.print
        input = self.input
        return self.outputs
    def __exit__(self, *args):
        global print
        global input
        del print
        del input
```

### Задача 1. Лучше гор могут быть только горы... (1 балл)

Напишите программу, которая запрашивает у пользователя размер обуви (целое число) и выводит на экран размер туфель для скалолазания, учитывая, что обычно размер скальных туфель на два размера меньше обычного.

Уточнение: считайте, что пользователь вводит допустимые значения размера обуви и знает, что мы ждем от него целые числа в качестве ответа. Задача не предполагает проверки ввода пользователя, но если интересно, можете над этим подумать.

**Пример работы программы**

Ввод:

    38

Вывод:

    36

<!-- grade_id: ans01, points: 0, grade: false, solution: true -->
```python
def problem01():

    # YOUR CODE HERE
    a = int(input())
    print(a-2)
problem01()
```

Output:
```
36
34
```

<!-- grade_id: test01, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("38", "36"),
    ("40", "38"),
    ("42", "40"),
    ("36", "34")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        problem01()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 2. В синем море, в белой пене... (1 балл)

В Светлогорске, в музее Мирового океана есть прекрасные весы, которые позволяют узнать свой вес, измеренный в селёдках, в китах, в креветках и в других морских обитателях.

Напишите программу, которая запрашивает у пользователя с клавиатуры его вес в килограммах (целое или дробное число) и выводит на экран его вес в селёдках, округлённый до целого числа. Считайте, что средний вес селёдки равен 350 граммам.

Уточнение: считайте, что пользователь вводит допустимые значения – целые и дробные числа больше 0. Задача не предполагает проверки ввода пользователя, но если интересно, можете над этим подумать.

**Пример работы программы**

Ввод:

    56

Вывод:

    Ваш вес в селёдках: 160.

<!-- grade_id: ans02, points: 0, grade: false, solution: true -->
```python
def problem02():

    # YOUR CODE HERE
    a = float(input())
    print('Ваш вес в селёдках: ',round(a/0.35),'.', sep='')
problem02()
```

Output:
```
64
Ваш вес в селёдках: 183.
```

<!-- grade_id: test02, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("56", "Ваш вес в селёдках: 160."),
    ("35", "Ваш вес в селёдках: 100."),
    ("64", "Ваш вес в селёдках: 183."),
    ("82", "Ваш вес в селёдках: 234."),
    ("40.5", "Ваш вес в селёдках: 116."),
    ("47.25", "Ваш вес в селёдках: 135.")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        problem02()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 3. Точность – вежливость королей (1 балл)

В машинном обучении существуют специальные меры, которые позволяют определить качество модели. Среди них — точность (*precision*), полнота (*recall*) и F-мера. F-мера вычисляется так:

$$
F = 2 \times \frac{\text{precision}\times \text{recall}}{\text{precision}+\text{recall}}
$$

Напишите программу, которая запрашивает с клавиатуры значение *precision*, затем значение *recall* (отдельный запрос для каждого значения), вычисляет F-меру и выводит её на экран с точностью до третьего знака после запятой.

Уточнение: считайте, что пользователь вводит допустимые значения – числа от 0 до 1 (оба числа не могут быть 0, иначе происходит деление на 0). Задача не предполагает проверки ввода пользователя, но если интересно, можете над этим подумать.

**Пример работы программы**

Ввод:

    0.9
    0.85

Вывод:

    F: 0.874

<!-- grade_id: ans03, points: 0, grade: false, solution: true -->
```python
def problem03():

    # YOUR CODE HERE
    precision=float(input())
    recall=float(input())
    print('F:',f'{round(2*precision*recall/(precision+recall),3):.3f}')

problem03()
```

Output:
```
0
1
F: 0.000
```

<!-- grade_id: test03, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("0.9 0.85", "F: 0.874"),
    ("0.65 0.68", "F: 0.665"),
    ("0.727 0.563", "F: 0.635"),
    ("0.41 0.5", "F: 0.451"),
    ("0 1", "F: 0.000"),
    ("1 0", "F: 0.000")
]

for inp, out in test_data:
    with Tester(inp.split()) as t:
        problem03()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 4. Power of numbers (1 балл)

Напишите программу, которая запрашивает с клавиатуры два целых положительных числа `a` и `b` (отдельный запрос для каждого значения) и выводит на экран строку следующего вида, с подставленными значениями на месте квадратных скобок:

    [a] в степени [b] равно [значение 1], [b] в степени [a] равно [значение 2]
    
Уточнение: считайте, что пользователь вводит допустимые значения – целые положительные числа. Задача не предполагает проверки ввода пользователя, но если интересно, можете над этим подумать.

**Примеры работы программы**

Ввод:

    2
    3

Вывод:

    2 в степени 3 равно 8, 3 в степени 2 равно 9

Ввод:

    4
    3

Вывод:

    4 в степени 3 равно 64, 3 в степени 4 равно 81

<!-- grade_id: ans04, points: 0, grade: false, solution: true -->
```python
def problem04():

    # YOUR CODE HERE
    a = int(input())
    b= int(input())
    print(a,' в степени ', b,' равно ',a**b,', ',b,' в степени ',a,' равно ',b**a, sep='')
problem04()
```

Output:
```
4
3
4 в степени 3 равно 64, 3 в степени 4 равно 81
```

<!-- grade_id: test04, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("4 3", "4 в степени 3 равно 64, 3 в степени 4 равно 81"),
    ("2 5", "2 в степени 5 равно 32, 5 в степени 2 равно 25"),
    ("3 6", "3 в степени 6 равно 729, 6 в степени 3 равно 216"),
    ("1 5", "1 в степени 5 равно 1, 5 в степени 1 равно 5"),
    ("1 1", "1 в степени 1 равно 1, 1 в степени 1 равно 1")
]

for inp, out in test_data:
    with Tester(inp.split()) as t:
        problem04()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 5. Ты не сорвёшь цветка, чтоб не поблёк он... (2 балла)

Согласно либретто оперы Ш.Гуно «Фауст», Мефистофель предсказал Зибелю, что он не сможет сорвать цветка, чтобы тот не завял. Итак, Зибель загадал следующее:

* Если он сорвёт две ромашки и в сумме число лепестков будет нечётным, он сможет собрать букет, который не завянет, и преподнести его Маргарите.

* Если он сорвёт две ромашки и в сумме число лепестков будет чётным, он попробует сорвать ещё третью. Если и тогда сумма будет чётной, он повременит с букетом и просто пойдёт грустить и думать дальше. А если сумма лепестков трёх ромашек всё же окажется нечётной, он попробует собрать букет.

Напишите программу, которая запрашивает с клавиатуры число лепестков на первой и второй ромашке (отдельный запрос для каждого значения) и выводит на экран количество ромашек, которое необходимо ещё сорвать для принятия окончательного решения.

При решении этой задачи нельзя использовать конструкцию `if-else` (даже если вы с ней знакомы) и любые способы явной проверки выполнения условий, например, с помощью оператора `==`.


**Пример работы программы**

Ввод:
    
    16
    17
    

Вывод:

    0
    
Ввод:
    
    14
    12
    

Вывод:

    1
    

Ввод:
    
    11
    19
    

Вывод:

    1    

<!-- grade_id: ans05, points: 0, grade: false, solution: true -->
```python
def problem05():

    # YOUR CODE HERE
    a = int(input())
    b = int(input())
    print((a+b+1)%2)

problem05()
```

Output:
```
16
17
0
```

<!-- grade_id: test05, points: 2, grade: true, solution: false -->
```python
test_data = [
    ("13 10", "0"),
    ("11 19", "1"),
    ("16 17", "0"),
    ("15 20", "0"),
    ("14 18", "1")
]

for inp, out in test_data:
    with Tester(inp.split()) as t:
        problem05()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 6.  Разделяй и властвуй (2 балла)

Напишите программу, которая запрашивает с клавиатуры целое положительное четырёхзначное число, сохраняет его в числовом формате (тип *integer*) и выводит на экран его представление в следующем виде (с подставленными значениями на месте квадратных скобок):

    [тысячи] * 1000 + [сотни] * 100 + [десятки] * 10 + [единицы]
    
Важно: «сохраняет в числовом виде» – обязательное условие, нельзя просто извлечь первый/второй/третий/четвёртый символ строки и подставить в текстовый шаблон, нужно работать с числом! Иначе это была бы простая задача :)

**Примеры работы программы**

Ввод:

    1234

Вывод:

    1 * 1000 + 2 * 100 + 3 * 10 + 4

Ввод:

    9000

Вывод:

    9 * 1000 + 0 * 100 + 0 * 10 + 0

<!-- grade_id: ans06, points: 0, grade: false, solution: true -->
```python
def problem06():

    # YOUR CODE HERE
    a = int(input())
    print(a//1000,' * 1000 + ',a//100-(a//1000)*10,' * 100 + ',a//10-(a//100)*10,' * 10 + ',a-(a//10)*10,sep='')
problem06()
```

Output:
```
1234
1 * 1000 + 2 * 100 + 3 * 10 + 4
```

<!-- grade_id: test06, points: 2, grade: true, solution: false -->
```python
test_data = [
    ("1234", "1 * 1000 + 2 * 100 + 3 * 10 + 4"),
    ("1046", "1 * 1000 + 0 * 100 + 4 * 10 + 6"),
    ("3600", "3 * 1000 + 6 * 100 + 0 * 10 + 0"),
    ("9701", "9 * 1000 + 7 * 100 + 0 * 10 + 1"),
    ("1000", "1 * 1000 + 0 * 100 + 0 * 10 + 0"),
    ("9001", "9 * 1000 + 0 * 100 + 0 * 10 + 1"),
    ("4280", "4 * 1000 + 2 * 100 + 8 * 10 + 0")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        problem06()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

### Задача 7.  Что скрывают звёзды (3 балла)

Напишите программу, которая угадывает, какая цифра была скрыта за «звёздочкой», если известно, что пользователь вводит с клавиатуры сначала одно двузначное число со «звёздочкой» на конце, затем – второе двузначное число со «звёздочкой» в начале, а затем их сумму. Считайте, что сумма чисел тоже является двузначным числом.

При решении этой задачи нельзя использовать конструкцию `if-else` (даже если вы с ней знакомы) и любые способы явной проверки выполнения условий, например, с помощью операторов `>` или `<`.

*Подсказка.* Извлечь первый элемент строки можно по индексу 0, а второй – по индексу 1, например, так:

    s = "abc"
    one = s[0]
    two = s[1]
    print(one)
    print(two)
    
    a
    b

**Примеры работы программы**

Ввод:

    2*
    *1
    54

Вывод:

    * = 3

Ввод:

    1*
    *3
    90

Вывод:

    * = 7


<!-- grade_id: ans07, points: 0, grade: false, solution: true -->
```python
def problem07():

    # YOUR CODE HERE
    a = str(input())
    b = str(input())
    c = int(input())
    k=c-int(a[0])*10-int(b[1])
    print('* =',k//10)
problem07()
```

Output:
```
1*
*3
90
* = 7
```

<!-- grade_id: test07, points: 3, grade: true, solution: false -->
```python
test_data = [
    ("2* *1 54", "* = 3"),
    ("1* *3 90", "* = 7"),
    ("2* *1 65", "* = 4"),
    ("3* *3 44", "* = 1"),
    ("1* *4 80", "* = 6"),
    ("1* *6 71", "* = 5")
]

for inp, out in test_data:
    with Tester(inp.split()) as t:
        problem07()
        assert len(t) == 1, "Вам нужно вывести ровно одну строку с ответом"
        assert t[0] == out+"\n", "Неверный ответ, была введена строка " + inp
```

*Названия и сюжеты задач:*
    
*№1: В.Высоцкий. Прощание с горами.* <br>
*№2: В синем море, в белой пене... (Арменфильм, 1984)* <br>
*№3: Калька с цитаты Людовика XIV* <br>
*№5: Опера «Фауст» (Ш.Гуно, 1859-1869) и фильм-опера «Фауст» (Б.Небиеридзе, 1982)*