# Программирование для всех <br>(основы работы с Python)

*Автор задач: Тамбовцева А.А.*

## Домашнее задание №4

**За разные задачи можно получить разное число баллов. Если не указано обратное, задача весит 1 балл. Чтобы получить оценку 10 за это задание, нужно набрать 10 баллов, то есть верно выполнить все задания.**

Почти все задачи снабжены тестами – ячейками со строками, содержащими конструкцию `assert`. Алгоритм работы с тестами такой: пишете код с решением, запускаете его, затем запускаете ячейку с тестами. Если после запуска ячейки с тестами ничего не произошло, ваш код с решением верный, если на экране появилось сообщение об ошибке, код неверный.

Во всех заданиях необходимо называть объекты так, как указано в условии, иначе решение будет считаться ошибочным (например, для Python вообще и для `assert` в частности объекты `n` и `N` – не одно и то же).

### Формат сдачи

Для сдачи домашнего задания необходимо скачать текущий ipynb-файл, вписать в него свои решения, сохранить изменения и загрузить обновленный файл на Dropbox по предложенной на странице курса ссылке. Иметь аккаунт Dropbox для этого не требуется, достаточно пройти по ссылке и выбрать файл на компьютере.

Получить ipynb-файл можно, выбрав в Jupyter пункт меню *File → Download as... → Notebook (.ipynb)*.

### Данные

**Файл для работы**

В файле `polit.csv` хранятся некоторые данные, используемые в [исследовании](http://jmcguire.faculty.wesleyan.edu/welcome/data/#20LActries) James W.McGuire, “Authoritarianism in the Middle East”.

**Показатели:**

* `ctry` – страна;
* `fh09` – индекс *Freedom House* за 2009 год (от 1 до 7, большие значения соответствуют меньшему уровню гражданских свобод);
* `polity09` – индекс *Polity2* за 2009 год (от -10 до 10, большие значения соответствуют большему уровню демократии);
* `gini` – индекс Джини за 2000-2010 годы (индекс имущественного неравенства);
* `fparl08` – доля женщин в парламенте в 2008 году (в процентах);
* `mena` – регион *Middle East and North Africa (1 – да, 0 – нет)*;
* `lati` – регион *Latin America*;
* `cari` – регион *Caribbean, former British, French, Dutch colony*;
* `east` – регион *East Asia*;
* `sovi` – регион *Former Soviet Bloc*;
* `afri` – регион *Africa*;
* `corr0509` – индекс *Control of Corruption*, среднее за 2005-2009 годы.

### Задание 1

Загрузите таблицу из файла `polit.csv` и сохраните датафрейм в переменную `polit`. Если в датафрейме есть строки с пропущенными значениями, удалите их. Сохраните изменения в `polit`.

**Подсказка:** удаление строк с пропущенными значениями – метод `.dropna()`.

<!-- grade_id: cell-bd09069837828d5c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import pandas as pd


file = pd.read_csv("polit.csv")
polit = file.dropna()
```

<!-- grade_id: cell-dc3ca2ea24ead192, points: 1, grade: true, solution: false -->
```python
assert polit.shape == (135, 13), "Проверьте, сохранились ли изменения в polit"
```

### Задание 2

Выберите в датафрейме строки, которые соответствуют странам со значениями индекса *Freedom House* (`fh09`) выше 5. Сохраните их в датафрейм `not_free`.

<!-- grade_id: cell-ef43df4f0d5c0bd2, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
not_free = polit[polit['fh09']>5]
```

<!-- grade_id: cell-515d575403b98e13, points: 1, grade: true, solution: false -->
```python
assert not_free.shape == (30, 13), "Неверное число строк или столбцов"
assert not_free.iloc[3, 9] == 0, "Неверный датафрейм"
assert not_free.iloc[19, 1] == 'Mauritania', "Неверный датафрейм"
```

### Задание 3

Выберите в датафрейме строки, которые соответствуют странам Африки (`afri`) с процентом женщин в парламенте (`fparl08`) выше 30%. Сохраните их в датафрейм `af_w`.

<!-- grade_id: cell-821b3c361b5f0e7d, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
polit.head()
af_w = polit[(polit['afri'] == 1) & (polit['fparl08']>30)]
```

<!-- grade_id: cell-4da7712178596899, points: 1, grade: true, solution: false -->
```python
assert af_w.shape == (6, 13), "Неверное число строк или столбцов"
assert af_w.iloc[3, 5] == 33.92, "Неверный датафрейм"
assert af_w.iloc[5, 7] == 0, "Неверный датафрейм"
```

### Задание 4

Выберите в датафрейме строки, которые соответствуют странам Африки или Латинской Америки (`afri` и `lati`) со значением индекса *Polity2* (`polity09`) больше или равным 8. Сохраните их в датафрейм `la_dem`.

<!-- grade_id: cell-505bba8952d6425e, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
la_dem = polit[((polit['afri'] == 1) | (polit['lati'] == 1))& (polit['polity09']>=8)]
```

<!-- grade_id: cell-451bfc8d30610712, points: 1, grade: true, solution: false -->
```python
assert la_dem.shape == (18, 13), "Неверное число строк или столбцов"
assert la_dem.iloc[3, 5] == 12.66, "Неверный датафрейм"
assert la_dem.iloc[5, 7] == 1, "Неверный датафрейм"
```

### Задание 5

Добавьте в датафрейм `polit` столбец `corr_round`, в котором будут храниться округленные до 2 знака после запятой значения индекса *Control of Corruption* (`corr0509`).

<!-- grade_id: cell-b773c30a03303cd1, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
polit['corr_round'] = polit['corr0509'].round(2)
```

Output:
```
<ipython-input-51-19207731be4c>:2: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  polit['corr_round'] = polit['corr0509'].round(2)
```

<!-- grade_id: cell-8c2782bc419be172, points: 1, grade: true, solution: false -->
```python
assert round(polit["corr_round"].sum()) == -14.0, "Ошибка в столбце corr_round"
assert polit["corr_round"].max() == 2.38, "Ошибка в столбце corr_round"
```

### Задание 6 (2 балла)

Добавьте в датафрейм `polit` столбец `fh_status`, в котором будут храниться типы стран в зависимости от значения индекса *Freedom House* (значения *free*, *partly free*, *not free*). Соответствие значений `fh09` с типами стран см. [здесь](https://freedomhouse.org/sites/default/files/2020-02/Methodology_FIW_2016.pdf) (*Table 3* на последней странице).

<!-- grade_id: cell-6d3e86f6514ffe7c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
polit.head()

def new_rank(x):
  global y
  if x >= 1 and x <= 2.5:
    y = 'free'
  elif x >= 3 and x <= 5:
    y = 'partly free'
  elif x >= 5.5 and x <= 7:
    y = 'not free'
  return y


polit['fh_status'] = polit["fh09"].apply(new_rank)
```

Output:
```
<ipython-input-53-7412549b1017>:15: SettingWithCopyWarning: 
A value is trying to be set on a copy of a slice from a DataFrame.
Try using .loc[row_indexer,col_indexer] = value instead

See the caveats in the documentation: https://pandas.pydata.org/pandas-docs/stable/user_guide/indexing.html#returning-a-view-versus-a-copy
  polit['fh_status'] = polit["fh09"].apply(new_rank)
```

<!-- grade_id: cell-bf606937b0f2cb58, points: 2, grade: true, solution: false -->
```python
assert polit["fh_status"].value_counts().values[0] == 57, "Неверные значения в столбце"
assert polit["fh_status"].value_counts().values[1] == 48, "Неверные значения в столбце"
assert polit["fh_status"].value_counts().values[2] == 30, "Неверные значения в столбце"
assert polit["fh_status"].values[3] == 'free', "Неверные значения в столбце"
assert polit["fh_status"].values[4] == 'partly free', "Неверные значения в столбце"
assert polit["fh_status"].values[134] == 'not free', "Неверные значения в столбце"
```

### Задание 7

*(Формат решения может быть разный, задача автоматически не тестируется, баллы будут добавлены при ручной проверке)*

Сгруппируйте строки в таблице в соответствии со значениями столбца `fh_status`, полученного в предыдущей части и выведите минимальное, среднее и максимальное значение показателя `gini` (индекс Джини) по каждой группе.

```python
# YOUR CODE HERE
polit.groupby('fh_status').agg(['min','mean','max'], numeric_only = True)['gini']
```

Output:
```
<ipython-input-55-26d71667307b>:2: FutureWarning: ['ctry'] did not aggregate successfully. If any error is raised this will raise in a future version of pandas. Drop these columns/ops to avoid this warning.
  polit.groupby('fh_status').agg(['min','mean','max'], numeric_only = True)['gini']
              min       mean   max
fh_status                         
free         24.7  38.170175  74.3
not free     16.8  39.750000  58.6
partly free  29.8  43.931250  64.3
```

### Задание 8 (2 балла)

*(Формат решения может быть разный, задача автоматически не тестируется, баллы будут добавлены при ручной проверке)*

Сгруппируйте строки в таблице в соответствии со значениями столбца `fh_status` и запишите строки, относящиеся к разным группам, в отдельные csv-файлы.

```python
# YOUR CODE HERE
polit.groupby('fh_status')

for group, tab in polit.groupby("fh_status"):
    fname = group + ".csv"
    tab.to_csv(fname)
```