# Программирование для всех (основы работы с Python)

## Экзамен

* Время выполнения экзамена: 2 часа 40 минут.
* Выполненную работу нужно загрузить на Dropbox по [ссылке](https://www.dropbox.com/request/GcOlWf1xTrdrTmRyETDb).
* Для автоматической проверки заданий части 1 не забудьте запустить ячейку с тестами ниже.

### Активация тестов

```python
# Фабрика тестов для проверки программ, принимающих данные через input()

from collections import deque

class Tester(object):
    def __init__(self, inp):
        self.outputs = []
        self.inputs = deque(inp)
    def print(self, *args, sep = " ", end = "\n"):
        text = sep.join(map(str, args)) + end
        newlines = text.splitlines(keepends=True)
        if self.outputs and self.outputs[-1] and self.outputs[-1][-1] != "\n" and newlines:
            self.outputs[-1] += newlines[0]
            self.outputs.extend(newlines[1:])
        else:
            self.outputs.extend(newlines)

    def input(self, *args):
        assert self.inputs, "Вы пытаетесь считать больше элементов, чем предусмотрено условием"
        return self.inputs.popleft()
    def __enter__(self):
        global print
        global input
        print = self.print
        input = self.input
        return self.outputs
    def __exit__(self, *args):
        global print
        global input
        del print
        del input
```

## Часть 1: базовый Python (18 баллов)

### Задача 0. Эпиграф  (1 балл)

В списке `quotes` хранятся имена героев, которые встречались в разных задачах на курсе, и их цитаты (воодушевляющие и не очень).

Напишите программу, которая выводит на экран строку с:

* выбранной случайным образом цитатой, которую можно использовать в качестве эпиграфа к экзамену по Python;
* именем героя, которому цитата принадлежит, в скобках через пробел.

Чтобы случайным образом выбрать элемент списка, воспользуйтесь функцией `choice()` из модуля `random`:

    from random import choice
    choice([1, 2, 5, 7])
    5

**Пример вывода цитаты**

    Ненавижу змей! Ой, это кто-то другой, а змеи – они ничего... (Зигзаг Маккряк)
    
Автоматического теста в этой задаче нет, так как псевдослучайный выбор может работать по-разному в зависимости от того, каким образом вы будете извлекать значения из списка `quotes`.

<!-- grade_id: cell-883fc14a17de6758, points: 0, grade: false, solution: true -->
```python
from random import choice

quotes = [("Клод Фролло", "Молиться бесполезно!"),
          ("Зибель", "Надейся на меня!"),
          ("Зигзаг Маккряк", "Ненавижу змей! Ой, это кто-то другой, а змеи – они ничего..."),
          ("Жак Веснушка", "Мы с тобой молодцы!"),
          ("Северус Снейп", "Может, вы еще не заметили, но жизнь вообще несправедлива.")]

### YOUR CODE HERE ###

choice([1, 2, 5, 7])
```

Output:
```
1
```

### Задача 1.  Let the feast begin (1 балл)

Хагрид выращивает гигантские тыквы, чтобы украсить Большой зал Хогвартса к Хэллоуину.
Помимо обломков своей волшебной палочки ему нужен специальный увеличительный раствор.
Известно, что на 1 килограмм тыквы нужно 10 капель раствора.

Напишите программу, которая запрашивает с клавиатуры вес тыквы в килограммах (целое положительное число) и выводит на экран необходимое число капель увеличителя (целое положительное число).

**Примеры работы программы:**

Ввод:

    20


Вывод:

    200

<!-- grade_id: cell-8b78c8d6ea8ea7b0, points: 0, grade: false, solution: true -->
```python
def increase():
    ### YOUR CODE HERE ###
    a = int(input())
    print(a*10)

increase()
```

Output:
```
20
200
```

<!-- grade_id: cell-c8012cd198b56e0f, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("20", "200"),
    ("10", "100"),
    ("5", "50"),
    ("7", "70"),
    ("17", "170"),
    ("33", "330")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        increase()
        assert len(t) == 1, "Неверный ответ"
        assert t[0] == out+"\n", "Неверный ответ"
```

### Задача 2.  Rat race (1 балл)

Рон очень переживает, что с его крысой Коростой что-то не так, потому что она стала чихать. Он вычитал в книжке, что для здоровых крыс норма – чихать не более 30 раз в день.

Напишите программу, которая запрашивает с клавиатуры число чиханий и выводит на экран сообщение `OK`, если число чихов не больше 30, и сообщение `BAD`, если число чихов более 30.

**Примеры работы программы:**

Ввод:

    25


Вывод:

    OK
    
Ввод:

    30


Вывод:

    OK    
    
Ввод:

    40


Вывод:

    BAD   

<!-- grade_id: cell-df6491b3a24c76c8, points: 0, grade: false, solution: true -->
```python
def sneeze():
    ### YOUR CODE HERE ###
    a = int(input())
    if a<=30:
      print('OK')
    else:
      print('BAD')
sneeze()
```

Output:
```
30
OK
```

<!-- grade_id: cell-6cb3d8609bc06236, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("20", "OK"),
    ("25", "OK"),
    ("0", "OK"),
    ("37", "BAD"),
    ("30", "OK"),
    ("45", "BAD"),
    ("52", "BAD"),
]

for inp, out in test_data:
    with Tester([inp]) as t:
        sneeze()
        assert len(t) == 1, "Неверный ответ"
        assert t[0] == out+"\n", "Неверный ответ"
```

### Задача 3. Hogwarts,  Hogwarts (2 балла)

В словаре `names` сохранены пары «название на словенском» – «название на английском»:

    names = {
    "Bradavičarka" : "Hogwarts",
    "Drznvraan" : "Ravenclaw",
    "Spolzgad" : "Slytherin",
    "Gryfondom" : "Gryffindor",
    "Pihpuff" : "HufflePuff",
    "Mrlakenstein" : "Voldemort",
    "Klobuk Izbiruh" : "Sorting hat"
    }
    
Напишите программу, которая запрашивает с клавиатуры слово на английском (с большой буквы или с маленькой, большими буквами или маленькими) и выводит на экран его эквивалент на словенском. Если такого слова в словаре нет, на экран должно выводиться сообщение `No information`.

**Примеры работы программы:**

Ввод:

    Ravenclaw


Вывод:

    Drznvraan
    
Ввод:

    HOGWARTS


Вывод:

    Bradavičarka  
    
Ввод:

    Sorting Hat


Вывод:

    Klobuk Izbiruh
    
Ввод:

    Castle


Вывод:

    No information

<!-- grade_id: cell-da22e28ac53f5801, points: 0, grade: false, solution: true -->
```python
from IPython.core.shellapp import List
def ling():

    names = {
        "Bradavičarka" : "Hogwarts",
        "Drznvraan" : "Ravenclaw",
        "Spolzgad" : "Slytherin",
        "Gryfondom" : "Gryffindor",
        "Pihpuff" : "HufflePuff",
        "Mrlakenstein" : "Voldemort",
        "Klobuk Izbiruh" : "Sorting hat"
        }

    ### YOUR CODE HERE ###


    list = [x.capitalize() for x in names.values()]
    a = str(input()).capitalize()
    if a in list:
      for key in names:
        if names[key].capitalize() == a:
          print(key, sep='')
    else:
      print('No information')

ling()
```

Output:
```
HufflePuff
Pihpuff
```

<!-- grade_id: cell-b17ef34ea5805bc7, points: 2, grade: true, solution: false -->
```python
test_data = [
    ("Ravenclaw", "Drznvraan"),
    ("Hogwarts", "Bradavičarka"),
    ("Slytherin", "Spolzgad"),
    ("Gryffindor", "Gryfondom"),
    ("HufflePuff", "Pihpuff"),
    ("Voldemort", "Mrlakenstein"),
    ("Sorting hat", "Klobuk Izbiruh"),
    ("ravenclaw", "Drznvraan"),
    ("HOGWARTS", "Bradavičarka"),
    ("GryffindoR", "Gryfondom"),
    ("Sorting Hat", "Klobuk Izbiruh"),
    ("voldemort", "Mrlakenstein"),
    ("SORTING HAT", "Klobuk Izbiruh"),
    ("snitch", "No information"),
    ("broom", "No information"),
    ("Snape", "No information")
]

for inp, out in test_data:
    with Tester([inp]) as t:
        ling()
        assert len(t) == 1, "Неверный ответ"
        assert t[0] == out+"\n", "Неверный ответ"
```

### Задача 4. Our dear Hogwarts (2 балла)

В списках `one` и `two` записаны словенские слова, имеющие отношение к волшебной вселенной Дж.Роулинг.

В списке `ch` собраны слова, которые пишутся через `č`, в списке – слова, которые пишутся через `ž`, в списке `c` – слова, которые пишутся через `c` и в списке `z` – слова, которые пишутся через `z`. Для простоты предположим, что пользователь при вводе слова может ошибиться только в одной из этих букв (`č` вместо `c` и наоборот, `ž` вместо `z` и наоборот).

Напишите программу, которая запрашивает с клавиатуры некоторое слово и:

* если слово написано правильно, выводит на экран это же слово;
* если слово написано неправильно, выводит на экран его в исправленном виде.

**Примеры работы программы:**

Ввод:

    čarovnik

Вывод:

    čarovnik
    
Ввод:

    tocka

Вывод:

    točka
    
    
Ввод:

    zaba

Вывод:

    žaba
    
Ввод:

    žemljevid

Вывод:

    zemljevid

<!-- grade_id: cell-fa0903ce8c37023d, points: 0, grade: false, solution: true -->
```python
# кому любопытно: čarovnik – волшебник, točka – балл, žival - животное,
# polž – слизняк, žaba - жаба, iskalec - ловец, zasledovalec – охотник,
# branilec – защитник, tolkač – загонщик, brazgotino – шрам,
# zgodovina – история, zemljevid – карта

def correct():

    ch = ["čarovnik", "točka", "tolkač"]
    zh = ["žival", "polž", "žaba"]

    c = ["iskalec", "zasledovalec", "branilec"]
    z = ["brazgotino", "zgodovina", "zemljevid"]


    ### YOUR CODE HERE ###
    a = str(input())
    if a in ch or a in zh or a in c or a in z:
      print(a)
    else:
      if 'z' in a:
        print(a.replace('z','ž'))
      elif 'ž' in a:
        print(a.replace('ž','z',))
      elif 'č' in a:
        print(a.replace('č','c'))
      elif 'c' in a:
        print(a.replace('c','č'))
correct()
```

Output:
```
žemljevid
zemljevid
```

<!-- grade_id: cell-7f95290288a3b49f, points: 2, grade: true, solution: false -->
```python
test_data = [
    ("čarovnik", "čarovnik"),
    ("carovnik", "čarovnik"),
    ("tocka", "točka"),
    ("točka", "točka"),
    ("žival", "žival"),
    ("zival", "žival"),
    ("polž", "polž"),
    ("polz", "polž"),
    ("žaba", "žaba"),
    ("zaba", "žaba"),
    ("iskalec", "iskalec"),
    ("iskaleč", "iskalec"),
    ("branilec", "branilec"),
    ("branileč", "branilec"),
    ("zemljevid", "zemljevid"),
    ("žemljevid", "zemljevid"),

]

for inp, out in test_data:
    with Tester([inp]) as t:
        correct()
        assert len(t) == 1, "Неверный ответ"
        assert t[0] == out+"\n", "Неверный ответ"
```

### Задача 5. Bad omen (1 балл)

Профессор Трелони гадает по книге. Она пять раз открывает книгу на случайно выбранной странице и смотрит на номер страницы. Если все пять раз книга открывается на странице 13, она ужасается и идет пить успокаивающий настой, если ни разу – она продолжает гадать дальше и выбирает номер строки на странице, если иное сочетание – сомневается и идет за стеклянным шаром.

Напишите программу, которая 5 раз запрашивает номер страницы и:

* если все пять раз введено число 13, выводит на экран сообщение `Keep calm`;
* если ни один раз не было введено число 13, выводит на экран сообщение `Go on`;
* если число 13 было введено иное число раз, выводит на экран сообщение `Crystal Ball`.


**Примеры работы программы:**

Ввод:

    13
    13
    13
    13
    13

Вывод:

    Keep calm
    
Ввод:

    13
    10
    10
    11
    13

Вывод:

    Crystal Ball
    
Ввод:

    1
    14
    12
    19
    100

Вывод:

    Go on

<!-- grade_id: cell-56e1eb0d11b2e5ae, points: 0, grade: false, solution: true -->
```python
def bad_fortune():

    ### YOUR CODE HERE ###
    a = int(input())
    b = int(input())
    c = int(input())
    d = int(input())
    e = int(input())

    if a==13 and b==13 and c==13 and d==13 and e==13:
      print('Keep calm')
    elif a==13 or b==13 or c==13 or d==13 or e==13:
      print('Crystal Ball')
    else:
      print('Go on')
bad_fortune()
```

Output:
```
13
12
12
14
15
Crystal Ball
```

<!-- grade_id: cell-0fd4134bd9107093, points: 1, grade: true, solution: false -->
```python
test_data = [("13,13,13,13,13", ["Keep calm"]),
             ("13,10,12,12,100", ["Crystal Ball"]),
             ("1,14,12,19,100", ["Go on"]),
             ("100,140,12,19,100", ["Go on"]),
             ("111,24,17,133,100", ["Go on"]),
             ("13,10,13,13,100", ["Crystal Ball"]),
             ("13,13,13,13,100", ["Crystal Ball"]),
             ("13,10,98,13,200", ["Crystal Ball"]),
             ]
for inp, out in test_data:
    with Tester(inp.split(",")) as t:
        bad_fortune()
        line_t = "".join(t)
        t = line_t.splitlines()
        assert len(t) == len(out), "Неверный ответ"
        for l_test, l_out in zip(t, out):
            assert len(l_test.split()) == len(l_out.split()), "Неверный ответ"
            for el_test, el_out in zip(l_test.split(), l_out.split()):
                assert el_test == el_out, "Неверный ответ"
```

### Задача 6.  Fanstastic beasts (1 балл)

Хагрид, мечтающий о драконе, случайно забрел в нелегальный зоомагазин с экзотическими животными в Лютном переулке.
Хагриду нужно выполнить много дел по поручению Дамблдора, поэтому, как только он увидит любого дракона, он сразу его купит и убежит, не глядя на остальных животных.

Напишите программу, которая работает со списком продаваемых животных и выводит на экран число животных, на которых должен посмотреть Хагрид, прежде чем увидит дракона.

Определите функцию `count_beasts()`, которая:

* принимает на вход один аргумент – список продаваемых животных `animals` (список строк);
* возвращает число элементов списка до первого названия, содержащего слово «дракон» (целое число).

<!-- grade_id: cell-9351e580938e2654, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
def count_beasts(a):
  empty=[]
  if 'dragon' in a[0]:
    return 0
  else:
    for i in a:
      if 'dragon' in i:
        break
      else:
        empty.append(i)
    return len(empty)

```

```python
# списки, на которых можно потестить свою функцию

L1 = ["manticore", "firecrab", "dugbog", "acromantula", "blast-ended skrewt",
      "hebridean black (dragon)", "occamy", "peruvian vipertooth (dragon)"]

L2 = ["hebridean black (dragon)", "firecrab", "acromantula", "blast-ended skrewt",
      "occamy", "peruvian vipertooth (dragon)"]

L3 = ["kelpie", "firecrab", "acromantula", "ukrainian ironbelly (dragon)",
      "swedish short-snout (dragon)", "blast-ended skrewt"]

count_beasts(L1)
```

Output:
```
5
```

<!-- grade_id: cell-c81fb49aaee3e9a2, points: 1, grade: true, solution: false -->
```python
L1 = ["manticore", "firecrab", "dugbog", "acromantula", "blast-ended skrewt",
      "hebridean black (dragon)", "occamy", "peruvian vipertooth (dragon)"]

L2 = ["hebridean black (dragon)", "firecrab", "acromantula", "blast-ended skrewt",
      "occamy", "peruvian vipertooth (dragon)"]

L3 = ["kelpie", "firecrab", "acromantula", "ukrainian ironbelly (dragon)",
      "swedish short-snout (dragon)", "blast-ended skrewt"]

assert count_beasts(L1) == 5, "Неверно"
assert count_beasts(L2) == 0, "Неверно"
assert count_beasts(L3) == 3, "Неверно"
assert count_beasts(L1 + L2) == 5, "Неверно"
assert count_beasts(L3[0:2] + L2) == 2, "Неверно"
```

### Задача 7. Never tickle a sleeping dragon  (1 балл)

Фред и Джордж, обидевшись на своего одногруппника, подложили ему в сумку волшебное перо со сломанной функцией автокоррекции.
Поломка выражается в следующем: сочетание *snake fangs* (змеиные клыки) при написании заменяется на *Snape fangs*. Одногруппник этим пером написал контрольную работу по зельеварению и ничего не заметил.

Напишите программу, которая запрашивает с клавиатуры текст ответов на вопрос об ингредиентах зелья и выводит на экран число штрафных очков, которое будет получено, если известно, что профессор Снейп за каждую подобную замену будет начислять 5 штрафных очков.

**Примеры работы программы:**

Ввод:

    1 porcupine quills, Snape fangs, horned slugs, 2 Snape fangs, wolfsbane

Вывод:

    10
    
Ввод:

    1 porcupine quills, horned slugs, 2 wolfsbane ?

Вывод:

    0
    
Ввод:

    1 fanged geranium, Snape fangs, 2 Snape fangs, chizpurfle fangs, 3 Snape fangs, wolfsbane

Вывод:

    15

<!-- grade_id: cell-432e37de69f8e795, points: 0, grade: false, solution: true -->
```python
def never_tickle():

    ### YOUR CODE HERE ###
    a = input()
    j=0
    for i in a.split(','):
      if 'Snape fangs' in i:
        j+=5
    print(j)

never_tickle()
```

Output:
```
1 fanged geranium, Snape fangs, 2 Snape fangs, chizpurfle fangs, 3 Snape fangs, wolfsbane
15
```

<!-- grade_id: cell-0bc377427138a2f7, points: 1, grade: true, solution: false -->
```python
test_data = [
    ("1 porcupine quills, Snape fangs, horned slugs, 2 Snape fangs, wolfsbane", "10"),
    ("1 porcupine quills, horned slugs, 2 wolfsbane ?", "0"),
    ("1 fanged geranium, Snape fangs, 2 Snape fangs, chizpurfle fangs, 3 Snape fangs, wolfsbane", "15"),
    ("1 Snape fangs, Snape fangs, wolfsbane", "10"),
    ("Snape fangs, porcupine quills", "5"),
    ("Snape fangs, Snape fangs, Snape fangs, Snape fangs", "20"),
    ("salamander blood, lionfish spines, flobberworm mucus", "0"),
]

for inp, out in test_data:
    with Tester([inp]) as t:
        never_tickle()
        assert len(t) == 1, "Неверный ответ"
        assert t[0] == out+"\n", "Неверный ответ"
```

### Задача 8.  Subtle science (2 балла)

Реализуйте программу, которая обновляет словарь с ингредиентами зелья.

Определите функцию `add_item()`, которая принимает три аргумента:

* `recipy` — непустой словарь;
* `item` — название ингредиента (строка);
* `quantity` – необходимое количество ингредиента в граммах (целое число).

Если запись с ключом `item` уже есть в словаре, соответствующее ему количество заменяется на значение `quantity`. Если записи с таким ключом нет, она добавляется в словарь. В итоге функция `add_item` возвращает обновленный словарь `recipy`.

**Примеры работы программы:**

Ввод:

    add_item({"ginger" : 25,
          "springwater": 400,
          "snake fangs" : 50}, "snake fangs", 100)
          
Результат:

    {'ginger': 25, 'springwater': 400, 'snake fangs': 100}
  
Ввод:

    add_item({"ginger" : 25,
          "springwater": 400,
          "snake fangs" : 50}, "lichi", 60)
          
Результат:

    {'ginger': 25, 'springwater': 400, 'snake fangs': 50, 'lichi': 60}

<!-- grade_id: cell-46d5495ed07c8462, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
def add_item(a,b,c):
  a[b]=c
  return a

add_item({"ginger" : 25,
      "springwater": 400,
      "snake fangs" : 50}, "snake fangs", 100)

```

Output:
```
{'ginger': 25, 'springwater': 400, 'snake fangs': 100}
```

<!-- grade_id: cell-6e833df0cb2d6559, points: 2, grade: true, solution: false -->
```python
assert add_item({"ginger" : 25,
      "springwater": 400,
      "snake fangs" : 50}, "snake fangs", 100) == {'ginger': 25,
                                                   'springwater': 400,
                                                   'snake fangs': 100}, "Неверный ответ"
assert add_item({"ginger" : 25,
      "springwater": 400,
      "snake fangs" : 50}, "lichi", 60) == {'ginger': 25,
                                            'springwater': 400,
                                            'snake fangs': 50,
                                            'lichi': 60}, "Неверный ответ"
```

### Задача 9. To sum up (3 балла)

Необходимо определить, размах оценок за какую из форм контроля по курсу трансфигурации больше (размах вычисляется как разность между максимальным и минимальным значением).

**ТРЕБУЕМАЯ ФУНКЦИЯ**

* Функция `compare()`, которая принимает на вход словарь и две строки.
* Ключами словаря являются имена студентов, значениями словаря – словари с их оценками по трем формам контроля. Гарантируется, что все вложенные словари состоят из одинакового числа элементов.
* Первая строка – название одной формы контроля, вторая строка – название второй формы контроля.
* Функция должна возвращать строку с названием формы контроля, размах оценок за которую больше.


**Примеры работы функции:**

Ввод:

    grades = {
        "Гермиона Грейнджер" : {"эссе" : 100, "тест" : 98,
                               "практикум" : 100},
        "Гарри Поттер" : {"эссе" : 68, "тест" : 70, "практикум" : 82},
        "Рон Уизли" : {"эссе" : 56, "тест" : 60, "практикум" : 44},
        "Лаванда Браун" : {"эссе" : 52, "тест" : 48, "практикум" : 56},
        "Полумна Лавгуд" : {"эссе" : 85, "тест" : 72, "практикум" : 92},
    }
    
    compare(grades, "эссе", "тест")

Результат:

    'тест'
    
    
Ввод:

    grades = {
        "Гермиона Грейнджер" : {"эссе" : 100, "тест" : 98,
                               "практикум" : 100},
        "Гарри Поттер" : {"эссе" : 68, "тест" : 70, "практикум" : 82},
        "Рон Уизли" : {"эссе" : 56, "тест" : 60, "практикум" : 44},
        "Лаванда Браун" : {"эссе" : 52, "тест" : 48, "практикум" : 56},
        "Полумна Лавгуд" : {"эссе" : 85, "тест" : 72, "практикум" : 92},
    }

    compare(grades, "практикум", "тест")

Результат:

    'практикум'
    
Ввод:

    grades = {
        "Гермиона Грейнджер" : {"эссе" : 95, "тест" : 94,
                               "практикум" : 87},
        "Гарри Поттер" : {"эссе" : 75, "тест" : 40, "практикум" : 34},
        "Рон Уизли" : {"эссе" : 73, "тест" : 40, "практикум" : 14}
    }

    compare(grades, "эссе", "практикум")

Результат:

    'практикум'    

<!-- grade_id: cell-b7cff997172c0dd3, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
def compare(a,b,c):
  j=[]
  k=[]

  for i in a.values():
    if b in i:
      j.append(i[b])
    if c in i:
      k.append(i[c])
  var1=max(j)-min(j)
  var2=max(k)-min(k)

  if var1>var2:
   return b
  else:
   return c

```

<!-- grade_id: cell-5445629b66f5547e, points: 3, grade: true, solution: false -->
```python
# словари, на которых можно потестить свою функцию

D1 = {
    "Гермиона Грейнджер" : {"эссе" : 100, "тест" : 98, "практикум" : 100},
    "Гарри Поттер" : {"эссе" : 68, "тест" : 70, "практикум" : 82},
    "Рон Уизли" : {"эссе" : 56, "тест" : 60, "практикум" : 44},
    "Лаванда Браун" : {"эссе" : 52, "тест" : 48, "практикум" : 56},
    "Полумна Лавгуд" : {"эссе" : 85, "тест" : 72, "практикум" : 92}
}

D2 = {
    "Гермиона Грейнджер" : {"эссе" : 95, "тест" : 94,
                           "практикум" : 87},
    "Гарри Поттер" : {"эссе" : 75, "тест" : 40, "практикум" : 34},
    "Рон Уизли" : {"эссе" : 73, "тест" : 40, "практикум" : 14}
}

compare(D1,"эссе", "практикум")
```

Output:
```
'практикум'
```

```python
D1 = {
    "Гермиона Грейнджер" : {"эссе" : 100, "тест" : 98, "практикум" : 100},
    "Гарри Поттер" : {"эссе" : 68, "тест" : 70, "практикум" : 82},
    "Рон Уизли" : {"эссе" : 56, "тест" : 60, "практикум" : 44},
    "Лаванда Браун" : {"эссе" : 52, "тест" : 48, "практикум" : 56},
    "Полумна Лавгуд" : {"эссе" : 85, "тест" : 72, "практикум" : 92}
}

D2 = {
    "Гермиона Грейнджер" : {"эссе" : 95, "тест" : 94,
                           "практикум" : 87},
    "Гарри Поттер" : {"эссе" : 75, "тест" : 40, "практикум" : 34},
    "Рон Уизли" : {"эссе" : 73, "тест" : 40, "практикум" : 14}
}

assert compare(D1, "эссе", "тест") == "тест", "Неверный ответ"
assert compare(D1, "практикум", "тест") == "практикум", "Неверный ответ"
assert compare(D2, "эссе", "практикум") == "практикум", "Неверный ответ"
```

### Задача 10. Wild problem (3 балла)

Хагрид производит контрольное взвешивание своих магических существ. Помогите ему собрать информацию об изменениях веса существ.

**ФОРМАТ ВВОДА**

* Записи о взвешенных существах: через точку с запятой название зверя, вес зверя при взвешивании в прошлом месяце, вес зверя в текущем месяце. Оба значения веса – целые числа, измерены в одинаковых единицах измерения. Каждая запись вводится с новой строки. Название существа может содержать любые символы, кроме точки с запятой.
* Так как Хагрид немного рассеян, названия существ могут повторяться. То есть по одному и тому же существу могут быть разные записи, в которых название зверя одинаковое, а значения веса в текущем месяце – разные из-за погрешности весов.
* Когда записи заканчиваются, на отдельной строке вводится слово "FINITA".

**ФОРМАТ ВЫВОДА**

* Список пар (одна пара – кортеж), в которых на первом месте должно быть название существа (строка), а на втором – набранный существом за месяц вес (целое число). Пары в списке должны быть отсортированы по убыванию набранного веса. При этом все звери в списке должны быть уникальны: если при вводе возникали записи с одинаковым названием существа, учитывалась только первая из них.

**Примеры работы программы**

Ввод:

    niffler;200;230
    bowtruckle;5;7
    baby dragon;1000;1200
    niffler;200;225
    porlock;600;650
    porlock;600;660
    FINITA

Результат:

    [('baby dragon', 200), ('porlock', 50), ('niffler', 30),('bowtruckle', 2)]

Ввод:

    bowtruckle;5;9
    bowtruckle;5;8
    fairy;2;2
    firecrab;400;700
    FINITA

Результат:

    [('firecrab', 300), ('bowtruckle', 4), ('fairy', 0)]

<!-- grade_id: cell-89a5241438aacb83, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
```

<!-- grade_id: cell-b67f4b4f8a73ea03, points: 3, grade: true, solution: false -->
```python
test_data = [("bowtruckle;5;9&bowtruckle;5;8&fairy;2;2&firecrab;400;700&FINITA",
              ["[('firecrab', 300), ('bowtruckle', 4), ('fairy', 0)]"]),
             ("niffler;200;230&bowtruckle;5;7&baby dragon;1000;1200&niffler;200;225&porlock;600;650&porlock;600;660&FINITA",
              ["[('baby dragon', 200), ('porlock', 50), ('niffler', 30), ('bowtruckle', 2)]"]),
             ("niffler;200;230&fairy;2;2&FINITA",
              ["[('niffler', 30), ('fairy', 0)]"])
             ]
for inp, out in test_data:
    with Tester(inp.split("&")) as t:
        problem16()
        line_t = "".join(t)
        t = line_t.splitlines()
        assert len(t) == len(out), "Неверный ответ"
        for l_test, l_out in zip(t, out):
            assert len(l_test.split()) == len(l_out.split()), "Неверный ответ"
            for el_test, el_out in zip(l_test.split(), l_out.split()):
                assert el_test == el_out, "Неверный ответ"
```

## Часть 2: библиотека pandas (12 баллов)

Так как задания в этой части либо простые, легко проверяемые визуально (можно посмотреть и увидеть, правильно ли что-то отфильтровано или добавлено), либо не предполагающие единственного верного формата решения, автоматических тестов тут нет.

### Данные

В файле `hp_upd.csv` сохранена информация по героям из вселенной Дж.Роулинг:

* `Name`: имя героя;
* `Gender`: пол героя;
* `Job`: должность;
* `House`: факультет;
* `Wand`: характеристика палочки;
* `Patronus`: патронус;
* `Species`: вид;
* `Blood.status`: статус крови;
* `Hair.colour`: цвет волос;
* `Eye.colour`: цвет глаз;
* `Loyalty`: указание, на чьей стороне;
* `Skills`: навыки и особые умения;
* `Birth`: дата рождения;
* `Death`: дата смерти;
* `Wand.length`: длина палочки в дюймах;
* `Birth.year`: год рождения;
* `End.year`: год смерти;
* `Age`: примерный возраст на 2021 год.

**(1 балл)** Загрузите файл с данными и сохраните таблицу в виде датафрйма `pandas`.

<!-- grade_id: cell-05c78d11d7c1842f, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
import pandas as pd

data = pd.read_csv("hp_upd.csv")
```

### Задача 1 (1 балл)

Переименуйте столбец `Gender` в `Sex` и сохраните изменения в датафрейме.

<!-- grade_id: cell-f55715ed983cc85d, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
data = data.rename(columns={'Gender': 'Sex'})
data['Sex']
```

Output:
```
0        Male
1        Male
2      Female
3        Male
4        Male
        ...  
135    Female
136      Male
137      Male
138      Male
139      Male
Name: Sex, Length: 140, dtype: object
```

### Задача 2 (1 балл)

Выберите строки, соответствующие героям из Гриффиндора женского пола, и сохраните их в датафрейм `g_fem`.

<!-- grade_id: cell-b9b49aff45c093c7, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
g_fem=data[data['Sex']=='Female']
```

### Задача 3 (1 балл)

Определите (используя `pandas`, не на глаз) число строк в датафрейме `g_fem`.

<!-- grade_id: cell-13af304dac8ea4fb, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
g_fem.shape[0]
```

Output:
```
49
```

### Задача 4 (1 балл)

Вычислите среднюю длину волшебных палочек и округлите ее до целых. Сохраните ответ в переменную `m`.

<!-- grade_id: cell-abb1b827dd79c408, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
m = data['Wand.length'].mean().round(0)
m

```

Output:
```
11.0
```

### Задача 5 (1 балл)

Выберите строки, соответствующие героям с волшебными палочками длины выше среднего, и сохраните их в датафрейм `high`.

<!-- grade_id: cell-7875462e940e7808, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
high = data[data['Wand.length']>m]
```

### Задача 6 (2 балла)

Выведите таблицу частот для показателя `House` в датафрейме со строками, отобранными в предыдущей задаче.

<!-- grade_id: cell-50fc0e71bb44ad12, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
df1 = data.groupby(['House']).count()
df1
```

Output:
```
Unnamed: 0  Id  Name  Sex  Job  Wand  Patronus  \
House                                                                          
Beauxbatons Academy of Magic           3   3     3    3    2     3         3   
Durmstrang Institute                   1   1     1    1    1     1         1   
Gryffindor                            38  38    38   38   34    38        38   
Hufflepuff                            13  13    13   13   12    13        13   
Ravenclaw                             18  18    18   18   18    18        18   
Slytherin                             28  28    28   27   24    28        27   

                              Species  Blood.status  Hair.colour  Eye.colour  \
House                                                                          
Beauxbatons Academy of Magic        3             3            3           2   
Durmstrang Institute                1             1            1           1   
Gryffindor                         38            37           36     
```

### Задача 7 (1 балл)

Добавьте в исходный датафрейм столбец `Wand.cm` со значениями длины волшебных палочек в сантиметрах (1 дюйм равен 2.54 сантиметрам).

<!-- grade_id: cell-3db8284352187de8, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
data['Wand.cm']=data['Wand.length']*2.54
```

### Задача 8 (1 балл)

Добавьте в исходный датафрейм столбец `Alive` из 0 и 1, где 1 соответствует героям, оставшимся в живых (год в `End.year` не 2021), а 0 – не оставшимся в живых (год 2021 в `End.year`).

<!-- grade_id: cell-50861a27f72ee774, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
data['Alive'] = data["End.year"].apply(lambda x: 1 if x == 2021 else 0)
```

### Задача 9 (2 балла)

Сгруппируйте строки в исходном датафрейме по значениям в столбце `Alive` и выведите для каждой группы средний и максимальный возраст героев.

<!-- grade_id: cell-eb2c94c6ac3d110a, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
data.groupby("Alive").agg(["mean","max"], numeric_only = True)['Age']
```

Output:
```
<ipython-input-116-f2ce6f019657>:2: FutureWarning: ['Name', 'Sex', 'Job', 'House', 'Wand', 'Patronus', 'Species', 'Blood.status', 'Hair.colour', 'Eye.colour', 'Loyalty', 'Skills', 'Birth', 'Death'] did not aggregate successfully. If any error is raised this will raise in a future version of pandas. Drop these columns/ops to avoid this warning.
  data.groupby("Alive").agg(["mean","max"], numeric_only = True)['Age']
            mean    max
Alive                  
0      38.740741  116.0
1      52.034091  141.0
```

## Часть 3: парсинг HTML (10 баллов)

Квиддич как своеобразный вид спорта существует и в реальном мире, только на метлах там не летают. Популярен этот спорт среди студентов.

Итак, дана страница с клубами по квиддичу в Великобритании: https://quadballuk.org/clubs.

Вам необходимо с этой страницы забрать:

* названия клубов;
* соответствующие им университеты или организации;
* полученные клубами звания, то есть медали.

Код страницы для восприятия довольно тяжел (фрагменты стилевого кода вписаны прямо в HTML, это ужасно), хотя в чистом виде структуры у страницы несложная. Очень рекомендуется использовать поиск по странице, без него находить элементы среди стилевых настроек CSS неудобно.

Для упрощения жизни рекомендуемая стратегия поиска информации:

1. Определить, в каких разделах (с какими тэгами и атрибутами) хранится информация по каждому клубу. То есть, понять какими элементами на странице представлены «карточки» клубов с фото и подписями.
2. Собрать все «карточки» в список. В каждой карточке найти по тэгам/атрибутам названия клубов, названия университетов/организаций, полученные награды.

Если планируете собрать только часть информации, стратегия может быть любой, но вот если доберетесь до наград, то есть ссылок на картинки с медалями, описанный выше механизм очень поможет, потому что найти названия клубов и названия университетов можно совершенно независимо (их одинаковое количество), а вот с медалями так не получится (у кого-то они есть, у кого-то их нет, да и количество наград разное).

### Задача 1 (1 балл)

Используя возможности `requests` и `bs4`, сохраните в переменную `soup` объект типа `BeautifulSoup` с исходным кодом изучаемой страницы.

<!-- grade_id: cell-96005b9681c0fda2, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
```

### Задача 2 (2 балла)

Сохраните в список `names` названия клубов. Это должен быть список из 36 строк (объектов типа `string`).

Первые три элемента списка для определенности:

    'Bangor Broken Broomsticks',
    'Birmingham Badgers',
    'Bristol Quadball Club'

<!-- grade_id: cell-55674a874ddf0491, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
```

### Задача 3 (2 балла)

Сохраните в список `orgs` названия университетов/организаций, с которыми аффилированы клубы. Это должен быть список из 36 строк (объектов типа `string`).

Первые три элемента списка для определенности:

    'Bangor University',
    'Selly Oak Park, Birmingham',
    'Clifton Downs'

<!-- grade_id: cell-99e8f9ed6d9635ea, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
```

### Задача 4 (4 балла)

Сохраните в список `medals` награды, полученные клубами. Это должен быть список из 36 списков (среди них пустые, если клуб наград не получал).

Первые пять элементов списка для определенности:

    ['British Quadball Cup Bronze Medalists', 'Development Cup Bronze Medalists'],
    [],
    ['British Quadball Cup Silver Medalists', 'Southern Cup Bronze Medalists'],
    [],
    ['Development Cup Bronze Medalists'],
    ['Northern Cup Silver Medalists', 'British Quadball Cup Bronze Medalists'],

<!-- grade_id: cell-67ea43c661e096ce, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
```

### Задача 5 (1 балл)

Объедините полученные списки в словарь, а затем – в датафрейм pandas, чтобы получилась аккуратная табличка из 36 строк и 3 столбцов.

<!-- grade_id: cell-e12b0d0feabf57be, points: 0, grade: false, solution: true -->
```python
### YOUR CODE HERE ###
```